<div style="background:#222222; color:#ffffff; padding:20px">
<h1 align="center">Guided ML With The Iris Dataset</h1>

<h2 align="center" tyle="color:#01ff84" >Learning type | Activity type | Objective |</h2>
<h2 align="center">| Supervised | Multiclass classification | Identify a flower's class |</h2>


<div>

Contents:

1. Loading the data
2. Setting up supervised learning problem (selecting features)
3. Creating a first model
    - Creating train and test datasets
    - Normalizing train and test
    - Fitting and predicting
4. Evaluate the frist model predictions
5. Crossvalidation of the model
6. Creating an end to end ML pipeline
    - Train/Test Split
    - Normalize
    - Crossvalidations
    - Model
    - fitting and predicting

## Instructions with NBGrader removed

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.


Contents:
1. Loading the data
2. Setting up supervised learning problem (selecting features)
3. Creating a first model
    - Creating train and test datasets
    - Normalizing train and test
    - Fitting and predicting
4. Evaluate the frist model predictions
5. Crossvalidation of the model
6. Creating an end to end ML pipeline
    - Train/Test Split
    - Normalize
    - Crossvalidations
    - Model
    - fitting and predicting

## About the dataset

[Iris](https://archive.ics.uci.edu/ml/datasets/iris) is a well-known multiclass dataset. It contains 3 classes of flowers with 50 examples each. There are a total of 4 features for each flower.

![](./classic-datasets/images/Iris-versicolor-21_1.jpg)

## Package setups

1. Run the following two cells to initalize the required libraries. 

```python
#to debug package errors
import sys
sys.path
sys.executable

```

Output:
```
'C:\\Users\\Ellefson\\miniconda3\\envs\\strive-ai\\python.exe'
```

```python
# Import needed packages
# You may add or remove packages should you need them
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
from sklearn import datasets
from sklearn import model_selection
from sklearn.preprocessing import StandardScaler
from sklearn.model_selection import GridSearchCV
from sklearn import preprocessing
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split, cross_validate
from sklearn.pipeline import make_pipeline

# Display plots inline and change plot resolution to retina
%matplotlib inline
%config InlineBackend.figure_format = 'retina'
# Set Seaborn aesthetic parameters to defaults
sns.set()
```

```python

```

## Step 1: Loading the data

1. Load the iris dataset using ```datasets.load_iris()```
2. Investigate the data structure with ```.keys()```
3. Construct a dataframe from the dataset
4. Create a 'target' and a 'class' column that contains the target names and values
5. Display a random sample of the dataframe 

<!-- grade_id: cell-c92cca7755a29a2c, points: 0, grade: false, solution: true -->
```python
def load_data():
    #load the dataset
    #return the dataset
    # YOUR CODE HERE
    dataset = datasets.load_iris()
    return dataset
```

<!-- grade_id: cell-6955aa3f0eb9b484, points: 1, grade: true, solution: false -->
```python
assert load_data()['data'].shape == (150,4)
```

```python
dataset = load_data()
```

```python
keys = list(dataset)
keys
```

Output:
```
['data',
 'target',
 'frame',
 'target_names',
 'DESCR',
 'feature_names',
 'filename',
 'data_module']
```

```python
dataset
```

Output:
```
{'data': array([[5.1, 3.5, 1.4, 0.2],
        [4.9, 3. , 1.4, 0.2],
        [4.7, 3.2, 1.3, 0.2],
        [4.6, 3.1, 1.5, 0.2],
        [5. , 3.6, 1.4, 0.2],
        [5.4, 3.9, 1.7, 0.4],
        [4.6, 3.4, 1.4, 0.3],
        [5. , 3.4, 1.5, 0.2],
        [4.4, 2.9, 1.4, 0.2],
        [4.9, 3.1, 1.5, 0.1],
        [5.4, 3.7, 1.5, 0.2],
        [4.8, 3.4, 1.6, 0.2],
        [4.8, 3. , 1.4, 0.1],
        [4.3, 3. , 1.1, 0.1],
        [5.8, 4. , 1.2, 0.2],
        [5.7, 4.4, 1.5, 0.4],
        [5.4, 3.9, 1.3, 0.4],
        [5.1, 3.5, 1.4, 0.3],
        [5.7, 3.8, 1.7, 0.3],
        [5.1, 3.8, 1.5, 0.3],
        [5.4, 3.4, 1.7, 0.2],
        [5.1, 3.7, 1.5, 0.4],
        [4.6, 3.6, 1. , 0.2],
        [5.1, 3.3, 1.7, 0.5],
        [4.8, 3.4, 1.9, 0.2],
        [5. , 3. , 1.6, 0.2],
        [5. , 3.4, 1.6, 0.4],
        [5.2, 3.5, 1.5, 0.2],
        [5.2, 3.4, 1.4, 0.2],
        [4.7, 3.2, 1.6, 0.2],
        [4.8, 3.1, 1.6, 0.2],
        [5.4, 3.4, 1.5, 0.4],
        [5.2, 4.1, 1.5, 0.1],
  
```

<!-- grade_id: cell-054f7e75dfc37e15, points: 0, grade: false, solution: true -->
```python
def dataset_to_pandas():
    #put the dataset into a pandas DF using the feature names as columnsç
    #rename the column name so the dont include the '(cm)'
    #add 2 columns one with the target and another with the target_names
    # YOUR CODE HERE
    df = pd.DataFrame(dataset.data, columns=dataset.feature_names)
    df.columns = [column.replace("(cm)", "").strip() for column in df.columns]
    df["target"] = dataset.target
    df["class"] = dataset.target_names[dataset.target]
    return df
```

```python
dataset_to_pandas()
```

Output:
```
sepal length  sepal width  petal length  petal width  target      class
0             5.1          3.5           1.4          0.2       0     setosa
1             4.9          3.0           1.4          0.2       0     setosa
2             4.7          3.2           1.3          0.2       0     setosa
3             4.6          3.1           1.5          0.2       0     setosa
4             5.0          3.6           1.4          0.2       0     setosa
..            ...          ...           ...          ...     ...        ...
145           6.7          3.0           5.2          2.3       2  virginica
146           6.3          2.5           5.0          1.9       2  virginica
147           6.5          3.0           5.2          2.0       2  virginica
148           6.2          3.4           5.4          2.3       2  virginica
149           5.9          3.0           5.1          1.8       2  virginica

[150 rows x 6 columns]
```

<!-- grade_id: cell-b18a01895567bcda, points: 2, grade: true, solution: false -->
```python
df = dataset_to_pandas()
assert df['sepal length'].shape == (150,)
assert df['sepal width'].shape == (150,)
assert df['petal length'].shape == (150,)
assert df['petal width'].shape == (150,)
assert df['target'].shape == (150,)
assert df['class'].shape == (150,)

```

### Question
Find the X and y values we're looking for. Notice that y is categorical and thus, we could **one-hot encode it** if we are looking at **class** or we can just pick **target**. In order to one hot encode we have  to re-shape `y` it using the **.get_dummies** function. 

## For the purpose of this exercise, do not use hot encoding, go only for target but think about if you have to drop it somewhere or not...

```python
df_iris = dataset_to_pandas()
```

```python
df_iris
```

Output:
```
sepal length  sepal width  petal length  petal width  target      class
0             5.1          3.5           1.4          0.2       0     setosa
1             4.9          3.0           1.4          0.2       0     setosa
2             4.7          3.2           1.3          0.2       0     setosa
3             4.6          3.1           1.5          0.2       0     setosa
4             5.0          3.6           1.4          0.2       0     setosa
..            ...          ...           ...          ...     ...        ...
145           6.7          3.0           5.2          2.3       2  virginica
146           6.3          2.5           5.0          1.9       2  virginica
147           6.5          3.0           5.2          2.0       2  virginica
148           6.2          3.4           5.4          2.3       2  virginica
149           5.9          3.0           5.1          1.8       2  virginica

[150 rows x 6 columns]
```

<!-- grade_id: cell-fd162bfaf4394e79, points: 0, grade: false, solution: true -->
```python
# virginica close to versicolor BAD
#0 - setosa
#1 - versicolor
#2 - virginica 

# Same distance between all the classes GOOD
#[1,0,0] - setosa
#[0,1,0] - versicolor
#[0,0,1] - virginica 


def ohe():
    # YOUR CODE HERE
    return pd.get_dummies(df_iris, columns=["target"])
```

```python
ohe()
```

Output:
```
sepal length  sepal width  petal length  petal width      class  \
0             5.1          3.5           1.4          0.2     setosa   
1             4.9          3.0           1.4          0.2     setosa   
2             4.7          3.2           1.3          0.2     setosa   
3             4.6          3.1           1.5          0.2     setosa   
4             5.0          3.6           1.4          0.2     setosa   
..            ...          ...           ...          ...        ...   
145           6.7          3.0           5.2          2.3  virginica   
146           6.3          2.5           5.0          1.9  virginica   
147           6.5          3.0           5.2          2.0  virginica   
148           6.2          3.4           5.4          2.3  virginica   
149           5.9          3.0           5.1          1.8  virginica   

     target_0  target_1  target_2  
0           1         0         0  
1           1         0         0  
2           1         0         
```

<!-- grade_id: cell-423f24383553b71a, points: 1, grade: true, solution: false -->
```python
ohe_data = ohe()

assert ohe_data.shape == (150,8)
```

## Step 2: Setting up supervised learning problem (selecting features)

Feature selection is an essential step in improving a model's perfromance. In the first version of the model we will use the **'sepal length'** and **'sepal width'** as predicting features. Later we will see the effect of adding additional features.

1. Assign the values of the 'target' to Y as a numpy array
2. Assign the remaining feature values to X as a numpy array
3. Check the shape of X and Y. Check the first few values.
    - Can we confirm our X and Y are created correctly?

<!-- grade_id: cell-c0cf28fd18447a0b, points: 0, grade: false, solution: true -->
```python
def target_to_numpy():
    # YOUR CODE HERE
    return df_iris["target"].values

def data_to_numpy():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return df_iris[['sepal length','sepal width']].values
```

<!-- grade_id: cell-aaf03dfa38041beb, points: 1, grade: true, solution: false -->
```python
Y = target_to_numpy()
X = data_to_numpy()
assert isinstance(Y, np.ndarray)
assert isinstance(X, np.ndarray)
assert X.shape == (150,2)
```

```python
print(X.shape)
print(Y.shape)
X[:10]
```

Output:
```
(150, 2)
(150,)
array([[5.1, 3.5],
       [4.9, 3. ],
       [4.7, 3.2],
       [4.6, 3.1],
       [5. , 3.6],
       [5.4, 3.9],
       [4.6, 3.4],
       [5. , 3.4],
       [4.4, 2.9],
       [4.9, 3.1]])
```

## Step 3: Creating the first model

In lecture we learned about creating a train and test datasets, normalizing, and fitting a model. In this step we will see how to build a simple version of this.

We have to be careful when constructing our train and test datasets. First, when we create train and test datasets we have to be careful that we always have the same datapoints in each set. Otherwise our results won't be reproduceable or we might introduce a bias into our model.

We also need to be attentive to when we normalize the data. What would be the effect of normalizing the data (i.e. with StandardScaler to a range between 0 - 1) before we create our train and test sets? Effectively we would use information in the test set to structure the values in the training set and vice versa. Therefore normalizing train and test independently is the preferred method.

1. Create X_train, X_test, Y_train, Y_test using ```train_test_split()``` with an 80/20 train/test split. Look in the SKLearn documentation to understand how the function works.
    - Inspect the first few rows of X_train.
    - Run the cell a few times. Do the first few rows change?
    - What option can we use in ```train_test_split()``` to stop this from happening?
2. Normalize the train and test datasets with ```StandardScaler```
    - We can fit the transform with ```.fit()``` and ```.transform()``` to apply it. Look in the documentation for an esample of how to do this.
    - Does it make sense to normalize Y_train and Y_test?
3. Initalize a ```LogisticRegression()``` model and use the ```.fit()``` method to initalize the first model.
    - We will pass the X_train and Y_train variables to the ```.fit()``` method.
    - Once the model is fit, use the ```.predict()``` with the X_test and save the output as predictions.

<!-- grade_id: cell-586cc3c4a3d849fa, points: 0, grade: false, solution: true -->
```python
#split train and test data 80/20
#your code here
# YOUR CODE HERE
X_train, X_test, Y_train, Y_test = train_test_split(X, Y, test_size=0.2, random_state=42)

print(X_train.shape)
print(Y_train.shape)
print(X_test.shape)
print(Y_test.shape)
```

Output:
```
(120, 2)
(120,)
(30, 2)
(30,)
```

```python
X_train[:5]
```

Output:
```
array([[4.6, 3.6],
       [5.7, 4.4],
       [6.7, 3.1],
       [4.8, 3.4],
       [4.4, 3.2]])
```

<!-- grade_id: cell-4415454a211f5895, points: 1, grade: true, solution: false -->
```python
assert X_train.shape == (120,2)
assert Y_train.shape == (120,)
assert X_test.shape  == (30,2)
assert Y_test.shape  == (30,)
```

<!-- grade_id: cell-d94703390548f2df, points: 0, grade: false, solution: true -->
```python
#normalize the dataset
#your code here
#create and fit the scaler object on the training data
# YOUR CODE HERE
sc = StandardScaler()
X_train = sc.fit_transform(X_train)
X_test = sc.transform(X_test)
```

<!-- grade_id: cell-fd525045d90798a3, points: 2, grade: true, solution: false -->
```python
assert np.amin(X_train) >= -2.5
assert np.amax(X_train) <= 3.2
assert np.amin(X_test) >= -2
assert np.amin(X_test) <= 2.75
```

<!-- grade_id: cell-a91abe39c7261df3, points: 0, grade: false, solution: true -->
```python
#initalize and fit with Logistic Regression
prediction = 0
#your code here

#initalize the logistic regressor
#make predictions
# YOUR CODE HERE
logisticmodel = LogisticRegression()
logisticmodel.fit(X_train,Y_train)
predictions = logisticmodel.predict(X_test)
```

<!-- grade_id: cell-61e5bc16122b7220, points: 2, grade: true, solution: false -->
```python
assert predictions.shape == (30,)
```

## Step 4: Evaluate the frist model's predictions

We will learn more about how to evaluate the performance of a classifier in later lessons. For now we will use % accuracy as our metric. It is important to know that this metric only helps us understand the specific performance of our model and not, for example, where we can improve it, or where it already perfoms well.

1. Use ```.score()``` to evaluate the performance of our first model.

<!-- grade_id: cell-a19f9db86d767c57, points: 0, grade: false, solution: true -->
```python
score = 0
#evaluating the performace of our first model
#your code here
# YOUR CODE HERE
score = logisticmodel.score(X_test, Y_test)
print(score)
```

Output:
```
0.9
```

<!-- grade_id: cell-300a4e426c8fcfa1, points: 1, grade: true, solution: false -->
```python
assert score >=0.7
```

<!-- grade_id: cell-f7dbd3aa1419397f, points: 1, grade: true, solution: false -->
```python
assert score >=0.72
```

<!-- grade_id: cell-3b6954a5b2f244cd, points: 1, grade: true, solution: false -->
```python
assert score >=0.73
```

## Step 5: Crossvalidation of the model
Our first model achived ~90% accruacy. This is quite good. How do we know it is reproducable? If we run the model again and our performance is 85% which is correct? And what about improving our model? Can you think of one thing we can do to potentially improve the model?

#### Crossvalidation
Corssvalidation is when we create multiple X and Y datasets. On each dataset we train and fit the model. We then average the results and return a 'crossvalidated' accruacy.

1. Initalize a new version of the model you trained above with the same paramters.
2. Use ```cross_validate()``` to run the model with 5 crossvalidation folds. 

<!-- grade_id: cell-49d1242d4949806f, points: 0, grade: false, solution: true -->
```python
#model with cross validation
#your code here
#cross validate the training set
# YOUR CODE HERE
clf_cv = logisticmodel.fit(X_train,Y_train)
# raise NotImplementedError()
cv = cross_validate(clf_cv, X_train, Y_train, cv=5)
def print_scores(cv):
    #print out cross validation scores
    [print('Crossvalidation fold: {}  Accruacy: {}'.format(n, score)) for n, score in enumerate(cv['test_score'])]
    #print out the mean of the cross validation
    print('Mean train cross validation score {}'.format(cv['test_score'].mean()))
    
print_scores(cv)
```

Output:
```
Crossvalidation fold: 0  Accruacy: 0.7083333333333334
Crossvalidation fold: 1  Accruacy: 0.75
Crossvalidation fold: 2  Accruacy: 0.875
Crossvalidation fold: 3  Accruacy: 0.8333333333333334
Crossvalidation fold: 4  Accruacy: 0.75
Mean train cross validation score 0.7833333333333334
```

<!-- grade_id: cell-cbdb686462858d91, points: 4, grade: true, solution: false -->
```python
assert len(cv['test_score']) == 5
assert max(cv['test_score']) >= 0.85
assert min(cv['test_score']) >= 0.69
assert cv['test_score'].mean() >= 0.77
```

## Step 6: Creating an end to end ML pipeline
Congraulations you've trained, crossvalidated, predicted, and evaluated your frist classifier. Now that you understand the basic steps we will look at a way to combine all these steps together.

Before we go further think about what you would have to do if you wanted to change the model. Intalize a new model, change the vairables, redo the cross validation...etc. Seems like a lot. And when we have to change lots of code it is easy to make mistakes. And what if you wanted to try many models and see which one performed best? Or try changing many different features? How could you do it without writing each one out as we have?

The solution is to use SKLearn's pipeline class. A pipeline is an object that will execute the various steps in the machine learning process. We can choose what elements we want in the pipeline and those that we do not. Once setup, we can rapidly change models, or input data and have it return our results in an ordered way.


1. Initalize a scaler and a classifer object like we did previously.
2. Use the ```make_pipeline()``` function to construct a transofmraiton pipeline for the scaler and the classifier
3. Input the pipeline object to the cross_validator and evaluate with 5 folds.
4. Print out your results (hint: make a function for repetitve tasks like printing)

<!-- grade_id: cell-82685348a4b1d16a, points: 0, grade: false, solution: true -->
```python
#define the scaler
#define the classifier
#make the pipeline
#run the cross validation
#print results
scaler = 0
classifier = 0
pipe = 0
scores = 0
# YOUR CODE HERE
scaler = StandardScaler()
classifier = LogisticRegression()
pipe = make_pipeline(scaler, classifier)
scores = cross_validate(clf_cv, X_train, Y_train, cv=5)
# Calling print_scores
print_scores(scores)
# raise NotImplementedError()

```

Output:
```
Crossvalidation fold: 0  Accruacy: 0.7083333333333334
Crossvalidation fold: 1  Accruacy: 0.75
Crossvalidation fold: 2  Accruacy: 0.875
Crossvalidation fold: 3  Accruacy: 0.8333333333333334
Crossvalidation fold: 4  Accruacy: 0.75
Mean train cross validation score 0.7833333333333334
```

<!-- grade_id: cell-0f5c78a78336a3c0, points: 4, grade: true, solution: false -->
```python
assert type(pipe) == type(make_pipeline(scaler, classifier))
assert len(cv['test_score']) == 5
assert max(cv['test_score']) >= 0.83
assert min(cv['test_score']) >= 0.69
assert cv['test_score'].mean() >= 0.74
```

## Challenge Exercise

In this notebook we only used two features to predict the class of the flower. We also did not do any hypter parameter tuning. The challenge is to impove the prediction results. Some ideas we can try:
1. Add features to the input and run the cross validation pipeline
2. Investigate how to use ```GridSearchCV```, a powerful funtion that searches through hyperparmetrs and does cross validation.
    - Hint: Input the pipeline directly into GridSearchCV
3. Try a different models like RandomForest or SVM.

```python
sorted(pipe.get_params().keys())
```

Output:
```
['logisticregression',
 'logisticregression__C',
 'logisticregression__class_weight',
 'logisticregression__dual',
 'logisticregression__fit_intercept',
 'logisticregression__intercept_scaling',
 'logisticregression__l1_ratio',
 'logisticregression__max_iter',
 'logisticregression__multi_class',
 'logisticregression__n_jobs',
 'logisticregression__penalty',
 'logisticregression__random_state',
 'logisticregression__solver',
 'logisticregression__tol',
 'logisticregression__verbose',
 'logisticregression__warm_start',
 'memory',
 'standardscaler',
 'standardscaler__copy',
 'standardscaler__with_mean',
 'standardscaler__with_std',
 'steps',
 'verbose']
```

<!-- grade_id: cell-f1ba9880447e80ec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#selected_parameters = {'penalty' : ['l1', 'elasticnet', 'l2', 'none'],
#         'C' : [0.1, 1, 10, 5, 0.2, -0.1, -0.5]}

selected_params = {'logisticregression__penalty' : ['l1', 'elasticnet', 'l2', 'none'],
         'logisticregression__C' : [0.1, 1, 5, 0.5, 10],
                  'logisticregression__max_iter' : [0, 10, 50, 100]}
```

```python
selected_params
```

Output:
```
{'logisticregression__penalty': ['l1', 'elasticnet', 'l2', 'none'],
 'logisticregression__C': [0.1, 1, 5, 0.5, 10],
 'logisticregression__max_iter': [0, 10, 50, 100]}
```

<!-- grade_id: cell-71c06d508e0294b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
grid_cv = GridSearchCV(pipe, selected_params, verbose=True, cv=5)
```

<!-- grade_id: cell-230a4548c4b43eb7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
grid_cv.fit(X_train, Y_train)
```

Output:
```
Fitting 5 folds for each of 80 candidates, totalling 400 fits
C:\Users\Ellefson\miniconda3\envs\strive-ai\lib\site-packages\sklearn\linear_model\_logistic.py:818: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  extra_warning_msg=_LOGISTIC_SOLVER_CONVERGENCE_MSG,
C:\Users\Ellefson\miniconda3\envs\strive-ai\lib\site-packages\sklearn\linear_model\_logistic.py:818: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for al
```

<!-- grade_id: cell-be4475eea0463f76, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
print(grid_cv.cv_results_)
```

Output:
```
{'mean_fit_time': array([0.00099688, 0.00079741, 0.00279269, 0.00239358, 0.00059791,
       0.00059848, 0.00418882, 0.0037899 , 0.00059838, 0.00059819,
       0.00339088, 0.00897579, 0.00059795, 0.00039873, 0.00359058,
       0.00957389, 0.00059857, 0.00039892, 0.00219417, 0.00219359,
       0.00059838, 0.00039892, 0.00378985, 0.00398889, 0.00039907,
       0.00059824, 0.00438828, 0.00857739, 0.000598  , 0.00039892,
       0.00418863, 0.00757971, 0.00059862, 0.00039897, 0.00219388,
       0.00239658, 0.0005981 , 0.00059843, 0.00398917, 0.00398922,
       0.00059843, 0.00059838, 0.00498638, 0.00997286, 0.00059838,
       0.00039887, 0.00538564, 0.00897613, 0.00039887, 0.00059848,
       0.00219412, 0.0023942 , 0.00019956, 0.00039878, 0.00378966,
       0.00398927, 0.00059848, 0.00059814, 0.00418887, 0.00997291,
       0.00039887, 0.00079775, 0.00418844, 0.00897317, 0.00039907,
       0.00079803, 0.00219393, 0.00219431, 0.00039892, 0.00059853,
       0.00439014, 0.00398855, 0.00059848, 0
```

<!-- grade_id: cell-b31a03b36b3beb45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
```

```python
from sklearn.ensemble import RandomForestClassifier
```

```python
# Random Forest
scaler_R = StandardScaler()
classifier_R = RandomForestClassifier()
random_pipe = make_pipeline(scaler_R, classifier_R)
```

```python
sorted(random_pipe.get_params().keys())
```

Output:
```
['memory',
 'randomforestclassifier',
 'randomforestclassifier__bootstrap',
 'randomforestclassifier__ccp_alpha',
 'randomforestclassifier__class_weight',
 'randomforestclassifier__criterion',
 'randomforestclassifier__max_depth',
 'randomforestclassifier__max_features',
 'randomforestclassifier__max_leaf_nodes',
 'randomforestclassifier__max_samples',
 'randomforestclassifier__min_impurity_decrease',
 'randomforestclassifier__min_samples_leaf',
 'randomforestclassifier__min_samples_split',
 'randomforestclassifier__min_weight_fraction_leaf',
 'randomforestclassifier__n_estimators',
 'randomforestclassifier__n_jobs',
 'randomforestclassifier__oob_score',
 'randomforestclassifier__random_state',
 'randomforestclassifier__verbose',
 'randomforestclassifier__warm_start',
 'standardscaler',
 'standardscaler__copy',
 'standardscaler__with_mean',
 'standardscaler__with_std',
 'steps',
 'verbose']
```

```python
selected_params_R = {'randomforestclassifier__n_estimators' : [1, 100, 500, 1000],
         'randomforestclassifier__max_depth': [0, 1, 2]}
```

```python
grid_cv_R = GridSearchCV(random_pipe, selected_params_R, verbose=True, cv=5)
```

```python
grid_cv_R.fit(X_train, Y_train)
```

Output:
```
Fitting 5 folds for each of 12 candidates, totalling 60 fits
C:\Users\Ellefson\miniconda3\envs\strive-ai\lib\site-packages\sklearn\model_selection\_validation.py:372: FitFailedWarning: 
20 fits failed out of a total of 60.
The score on these train-test partitions for these parameters will be set to nan.
If these failures are not expected, you can try to debug them by setting error_score='raise'.

Below are more details about the failures:
--------------------------------------------------------------------------------
20 fits failed with the following error:
Traceback (most recent call last):
  File "C:\Users\Ellefson\miniconda3\envs\strive-ai\lib\site-packages\sklearn\model_selection\_validation.py", line 681, in _fit_and_score
    estimator.fit(X_train, y_train, **fit_params)
  File "C:\Users\Ellefson\miniconda3\envs\strive-ai\lib\site-packages\sklearn\pipeline.py", line 394, in fit
    self._final_estimator.fit(Xt, y, **fit_params_last_step)
  File "C:\Users\Ellefson\miniconda3\envs
```

```python
print(grid_cv_R.cv_results_)
```

Output:
```
{'mean_fit_time': array([0.00219378, 0.03091698, 0.14281754, 0.29401526, 0.00199447,
       0.09694095, 0.48051472, 0.98596296, 0.00199418, 0.10412159,
       0.48709731, 0.99892821]), 'std_fit_time': array([3.98827312e-04, 1.09262341e-03, 2.98516945e-03, 1.70859988e-02,
       5.13569337e-07, 2.47514585e-03, 1.20572338e-02, 1.41833178e-02,
       5.91739352e-07, 9.21864594e-03, 1.29363967e-02, 2.16458962e-02]), 'mean_score_time': array([0.        , 0.        , 0.        , 0.        , 0.00079789,
       0.00817771, 0.03809786, 0.08198071, 0.00019984, 0.00917549,
       0.04089026, 0.0779911 ]), 'std_score_time': array([0.        , 0.        , 0.        , 0.        , 0.00039895,
       0.00039902, 0.00074603, 0.00593086, 0.00039968, 0.00074662,
       0.00260007, 0.00116297]), 'param_randomforestclassifier__max_depth': masked_array(data=[0, 0, 0, 0, 1, 1, 1, 1, 2, 2, 2, 2],
             mask=[False, False, False, False, False, False, False, False,
                   False, False, False,
```

```python

```